# Дополнительное задание

### Библиотеки
Во всех задачах можно использовать библиотеки `requests` и `BeautifulSoup` (`bs4`).

### Задача 1
Написать функцию `get_news(url)`, принимающую на вход адрес веб-страницы `url`, загружающую эту веб-страницу и проверяющую, встречается ли в ней слово `Harry` (с большой буквы). Функция должна возвращать `True`, если встречается, и `False` в противном случае. Также функция должна возвращать `False`, если страницу не удалось открыть (например, была получена ошибка *404 Not Found*.)

<!-- grade_id: pr_1, points: 0, grade: false, solution: true -->
```python
from bs4 import BeautifulSoup        
import re
import requests
def get_news(url):
    try:
        page = requests.get(url,timeout=5)
        con = BeautifulSoup(page.content)
        
    except requests.exceptions.HTTPError:
        return False
    else:
        for i in con.find_all():
            if ("Harry" in i.text):
                return True
        return False

print (get_news("https://en.wikipedia.org/w/index.php?title=J._K._Rowling&oldid=694008857"))
print (get_news("https://en.wikipedia.org/w/index.php?title=Star_Wars&oldid=694701430"))
print (get_news("https://en.wikipedia.org/w/index.php?title=Darth_Vader&oldid=694617684"))
```

Output:
```
True
True
False
```

<!-- grade_id: pr_1_1, points: 1, grade: true, solution: false -->
```python
assert get_news("https://en.wikipedia.org/w/index.php?title=J._K._Rowling&oldid=694008857")
assert get_news("https://en.wikipedia.org/w/index.php?title=Star_Wars&oldid=694701430")
assert not get_news("https://en.wikipedia.org/w/index.php?title=Darth_Vader&oldid=694617684")
```

<!-- grade_id: pr_1_2, points: 1, grade: true, solution: false -->
```python
print (get_news("http://math-info.hse.ru/there_is_no_Harry_here"))
```

Output:
```
False
```

### Задача 2
Написать функцию `get_strong(html)`, принимающую на вход html-страницу в виде длинной строки, записанной в переменную `html`, и возвращающую строчку, содержащуюся в первом теге `strong`.

Примеры см. в тестах.

**Подсказка.** Вы можете создать объект `BeautifulSoup`.

<!-- grade_id: pr_2, points: 0, grade: false, solution: true -->
```python
from bs4 import BeautifulSoup
import re
def get_strong(html):
    con = BeautifulSoup(html)
    for i in con.find_all("strong"):
        return i.text
    
    
print (get_strong("<html><body><p>Hello, <strong>World</strong>!")) #"World"
html = """<html>
    <body>
        <p>
            Hello,
            <strong>
                World
            </strong>
        </p>
    </body>
</html>"""
    
print (get_strong(html).strip()) #"World"
print (get_strong("<html><body><p>tag &lt;strong&gt; is used in HTML\n to make letters <strong>stronger</strong>")) #"stronger"
print (get_strong("<html><body><strong>One\nTwo</strong></body></html>")) #"One\nTwo"
```

Output:
```
World
World
stronger
One
Two
```

<!-- grade_id: pr_2_2, points: 1, grade: true, solution: false -->
```python
assert get_strong("<html><body><p>Hello, <strong>World</strong>!") == "World"
html = """<html>
    <body>
        <p>
            Hello,
            <strong>
                World
            </strong>
        </p>
    </body>
</html>"""
assert get_strong(html).strip() == "World"
assert get_strong("<html><body><p>tag &lt;strong&gt; is used in HTML\n to make letters <strong>stronger</strong>") == "stronger"
assert get_strong("<html><body><strong>One\nTwo</strong></body></html>") == "One\nTwo"
```

### Задача 3
Для вставки картинок в HTML используется тег `<img>`, содержащий параметр `src` — адрес файла с картинкой. Например, `<img src="https://upload.wikimedia.org/wikipedia/commons/b/bd/Struthio_camelus_portrait_Whipsnade_Zoo.jpg"/>`. Написать функцию `all_images_src(html)`, принимающую на вход длинную строчку с HTML-документом, а возвращающую список адресов всех картинок, встречающихся в этом документе (в том порядке, в котором они встречаются в документе).

<!-- grade_id: pr_3, points: 0, grade: false, solution: true -->
```python
from bs4 import BeautifulSoup
import re
def all_images_src(html):
    cont = BeautifulSoup(html)
    array = cont.find_all("img", src = True)
    real_array = []
    for i in array:
        i = str(i)        
        i = re.search(r'src=(.+\.jpg)|(.+\.png)', i)   
        i = i.group(0)
        i=re.sub(".*=\"", "", i)
        real_array.append(i)
    return real_array


print(all_images_src('<html><body><img src="https://upload.wikimedia.org/wikipedia/commons/b/bd/Struthio_camelus_portrait_Whipsnade_Zoo.jpg"/>'))
print (all_images_src('<html><body><IMG src="test.jpg">\n'
                        '<p>Some text\n'
                        '<img SRC=\'well.png\'>\n'
                        '</p></body></html>'))
print (all_images_src('<html><body><p><a href="link.html">'
                      '<img alt="Just a test image" src="this is a test.jpg"><ul>' + "\n"
                      .join("<li><img src='img%04i.png'></li>" % i for i in range(1000)) + 
                      "</ul></p></body></html>"))

```

<!-- grade_id: pr_3_t, points: 1, grade: true, solution: false -->
```python
assert all_images_src('<html><body><img src="https://upload.wikimedia.org/wikipedia/commons/b/bd/Struthio_camelus_portrait_Whipsnade_Zoo.jpg"/>') == ["https://upload.wikimedia.org/wikipedia/commons/b/bd/Struthio_camelus_portrait_Whipsnade_Zoo.jpg"]
assert all_images_src( ('<html><body><IMG src="test.jpg">\n'
                        '<p>Some text\n'
                        '<img SRC=\'well.png\'>\n'
                        '</p></body></html>') ) == ["test.jpg", "well.png"]
assert all_images_src('<html><body><p><a href="link.html">'
                      '<img alt="Just a test image" src="this is a test.jpg"><ul>' + "\n"
                      .join("<li><img src='img%04i.png'></li>" % i for i in range(1000)) + 
                      "</ul></p></body></html>"
                     ) == ['this is a test.jpg'] + ['img%04i.png' % i for i in range(1000)]
```

<!-- grade_id: pr_4, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: pr_4_t, points: 2, grade: true, solution: false -->
```python

```

<!-- grade_id: pr_5, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: pr_5_1, points: 4, grade: true, solution: false -->
```python

```

```python

```